# Assignment 0: Tools and setting up your computer

This assignment simply checks that you can access our assignments and will familiarize you with submitting assignments via nbgrader.


**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INLCUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (by which we mean _delete_) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.

## Q1

Create a variable called `my_name`, and assign a string to it that is your name. In other words, enter `my_name = 'name'`. Be sure to put your name in single quotes, e.g., `'Brad'`.

<!-- grade_id: cell-4995bf24615505c5, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
my_name = 'string'
### END SOLUTION
```

<!-- grade_id: cell-8280dfa952e91135, points: 10, grade: true, solution: false -->
```python
## TESTS FOR Q1, worth 10 points
# This is a test cell. It checks that the code you wrote does what it's supposed to.
# You can use the information in this cell to make sure you've met the expectations for the question.
assert  my_name                     # This tests that a variable "my_name" exists
assert  isinstance(my_name, str)    # isinstance() tests that "my_name" is of a particular type (here, a string)
```

## Q2
Occasionally, we'll use these notebooks to ask multiple choice questions. When you respond to a multiple choice question, **be sure to put your response in single quotation marks, the way you did in Q1**. This will ensure that Python will read your response as a **string**.

Unless otherwise stated, the variable name for the answer to these kinds of questions should always be in the format of qX_answer, where the q means "**q**uestion" and the X is replaced by the question number. Since this is Q2, your variable name will therefore be `q2_answer`.

For example, if I wanted to respond with a hypothetical answer "E", I would write `q2_answer = 'E'`

Respond to the question below:

What is the **kernel**?

- A) A list of the variables stored in Python memory.
- B) An object that returns a string.
- C) A way to specify that you are using math variables and not programming variables.
- D) Where your code runs; your Jupyter notebook is connected to the kernel.

<!-- grade_id: cell-8d20cdb4d0f455ff, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q2_answer = 'D'
### END SOLUTION
```

<!-- grade_id: cell-1f9b7f3a324344ed, points: 10, grade: true, solution: false -->
```python
## TESTS FOR Q2, worth 10 points [Note: Includes A Hidden Test]
assert q2_answer

### BEGIN HIDDEN TESTS
assert q2_answer == 'D'
### END HIDDEN TESTS
```

## Q3
According to Professor Voytek, why is neuroanatomical knowledge especially important for *neural data science*?

- A) Neuroanatomical knowledge helps us create better encoding/decoding models.
- B) Neuroanatomy isn't important for neural data science.
- C) Neuroanatomy is critcial for patient diagnosis.
- D) Neuroanatomical knowledge helps research collaborations.

<!-- grade_id: cell-70e47f12b6d77d4c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q3_answer = 'A'
### END SOLUTION
```

<!-- grade_id: cell-71e6e1f3e847e016, points: 10, grade: true, solution: false -->
```python
## TESTS FOR Q3, worth 10 points [Note: Includes A Hidden Test]
assert q3_answer

### BEGIN HIDDEN TESTS
assert q3_answer == 'A'
### END HIDDEN TESTS
```

## Q4

Write code below that
1. Assigns the value "2" to `a`
2. Assigns the value "4" to `b`
3. Multiplies `a` times `b` and saves it as a variable, `c`.

<!-- grade_id: cell-3d608ef3ac4ec904, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
a = 2
b = 4
c = a * b
### END SOLUTION
```

<!-- grade_id: cell-9f60eddf4367b314, points: 15, grade: true, solution: false -->
```python
## TESTS FOR Q4, worth 15 points [Note: Hidden Tests.]
assert a and b and c

### BEGIN HIDDEN TESTS
assert a == 2
assert b == 4
assert c == 8
### END HIDDEN TESTS
```

## Q5
How are you feeling so far? Create a variable `feeling` and assign your response to it.

- A) No problems!
- B) This isn't so bad.
- C) I had a lot of trouble getting this to work.
- D) I'm deeply, terribly confused.

<!-- grade_id: cell-9e5ee257e7bd6c04, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
feeling = 'A'
### END SOLUTION
```

<!-- grade_id: cell-140a35f46e79f334, points: 5, grade: true, solution: false -->
```python
## TESTS FOR Q5, worth 5 points
assert feeling
```