# Assignment 1: Programming Fundamentals

This assignment covers programming fundamentals, to make sure people are all up to speed.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (by which we mean _delete_) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.


## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

#### Printing Variables
You can (and should) print and check variables as you go. This allows you to check what values they hold, and fix things if anything unexpected happens.

#### Restarting the Kernel
- If you run cells out of order, you can end up overwriting things in your namespace. 
- If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.
- Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.

# Part I: Defining & manipulating variables

## Q1
Create the following variables:
- a variable called `number` that stores the value 57
- a variable called `decimal` that stores the value 9.32
- a variable called `word` that stores the value 'COGS138'
- a variable called `int_dec`, created by casting `decimal` into an int

<!-- grade_id: cell-0ef9f3314e11ed42, points: 0, grade: false, solution: true -->
```python
# Create the variables

### BEGIN SOLUTION
number = 57
decimal = 9.32
word = 'COGS138'
int_dec = int(decimal)
### END SOLUTION
```

<!-- grade_id: cell-ad80ab07f7755947, points: 10, grade: true, solution: false -->
```python
## Tests for Q1, worth 10 points

# Check that variable `number` is an int, and has the correct value
assert isinstance(number, int)
assert number == 57

# Check that variable `decimal` is a float, and has the correct value
assert isinstance(decimal, float)
assert decimal == 9.32

# Check that variable `words` is a string, and has the correct value
assert isinstance(word, str)
assert word == "COGS138"

# Check that variable `int_dec is an int, and has the correct value
assert isinstance(int_dec, int)
assert int_dec == 9
```

## Q2

Create three different lists:

- A list called `string_list` that contains three strings (can be any strings)
- A list called `float_list` that contains three float numbers (can by any floats)
- A list called `int_list` that contains three integers (can be any integers)
- Then, create a new list, called `nested_list` which contains the three lists you created above.

<!-- grade_id: cell-951cc83732d3279e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
string_list = ['a', 'b', 'c']
float_list = [12.4, 16.4, 11.2]
int_list = [8, 2, 3]
nested_list = [string_list, float_list, int_list]
### END SOLUTION
```

<!-- grade_id: cell-b456dc57e71863e9, points: 10, grade: true, solution: false -->
```python
# Tests for Q2, worth 15 points total

assert isinstance(string_list, list)
assert len(string_list) == 3
assert isinstance(string_list[0], str)

assert isinstance(float_list, list)
assert len(float_list) == 3
assert isinstance(float_list[0], float)

assert isinstance(int_list, list)
assert len(int_list) == 3
assert isinstance(int_list[0], int)
```

<!-- grade_id: cell-95c2533a89d58338, points: 5, grade: true, solution: false -->
```python
# Tests for Q2
assert isinstance(nested_list, list)
assert len(nested_list) == 3
assert isinstance(nested_list[0], list)
```

# Part II: Indexing & slicing lists

For Part II, the following lists are provided for you. Make sure you run the following cell.

```python
list_1 = [10, 20, 30, 40]
list_2 = [13, 15, 17, 19, 21, 23]
list_3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11]
list_4 = [21, 9, 98, 289, 938]
```

## Q3

Do the following using indexing:

- Store the first value in `list_1` a variable called `index_1`
- Store the last value in `list_2` to a variable called `index_2`
- Store the first four values of `list_3` to a variable called `index_3`
- Store the last four values of `list_4` to a variable called `index_4`

<!-- grade_id: cell-eed94ec7e065c79f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
index_1 = list_1[0]
index_2 = list_2[-1]
index_3 = list_3[0:4]
index_4 = list_4[-4:]
### END SOLUTION
```

<!-- grade_id: cell-5bcb79cb6e723b00, points: 5, grade: true, solution: false -->
```python
# Tests for Q3, worth 20 points total.

assert isinstance(index_1, int)

### BEGIN HIDDEN TESTS
assert index_1 == 10
### END HIDDEN TESTS
```

<!-- grade_id: cell-484ae049c7ae1620, points: 5, grade: true, solution: false -->
```python
# Tests for Q3
assert isinstance(index_2, int)
### BEGIN HIDDEN TESTS
assert index_2 == 23
### END HIDDEN TESTS
```

<!-- grade_id: cell-f3b62401bfedf7ca, points: 5, grade: true, solution: false -->
```python
# Tests for Q3
assert isinstance(index_3, list)
assert len(index_3) == 4
### BEGIN HIDDEN TESTS
assert index_3 == [1, 2, 3, 4]
### END HIDDEN TESTS
```

<!-- grade_id: cell-8f453f7b2d2513fa, points: 5, grade: true, solution: false -->
```python
# Tests for Q3
assert isinstance(index_4, list)
assert len(index_4) == 4
### BEGIN HIDDEN TESTS
assert index_4 == [9, 98, 289, 938]
### END HIDDEN TESTS
```

# Part III: Create a gene dictionary
Here, we'll create a dictionary of genes that are expressed in the brain. Choose any of the genes below to put into your dictionary:

| Symbol | Full name                                |
|--------|------------------------------------------|
| BDNF   | brain derived neurotropic factor         |
| DISC1  | disrupted in schizophrenia 1             |
| PAX5   | paired box 5                             |
| MECP2  | methyl CpG binding protein 2             |
| SNAP25 | synaptosomal nerve-associated protein 25 |


## Q4
In the cell below, create a gene dictionary called `gene_dict` for three different genes from the list above. For each entry, the `key` should be a string that is the gene symbol, and the `value` should be the full name.

<!-- grade_id: cell-2328a4cfec50de9d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
gene_dict = {'Gene': 'boop', 'Gene2': 'boop boop', 'Gene3': 'boop boop boop'}
### END SOLUTION
```

<!-- grade_id: cell-8a3567cf1d345142, points: 10, grade: true, solution: false -->
```python
# Tests for Q4, worth 20 points total
assert isinstance(gene_dict, dict)
```

<!-- grade_id: cell-0e2562c1a5be1d9d, points: 10, grade: true, solution: false -->
```python
# Hidden tests for Q4 (no need to do anything here)
### BEGIN HIDDEN TESTS
# Check length of dictionary
assert len(gene_dict) == 3
# Check that entries are strings (probably unnecessary)
for value in gene_dict.values():
    print(value)
    assert isinstance(value,str)
### END HIDDEN TESTS
```