# Assignment 4: Correlations

This assignment will demonstrate the strengths of weaknesses of correlations, and especially inferring causation from them.

**PLEASE DO NOT CHANGE THE NAME OF THIS FILE.**

**PLEASE DO NOT COPY & PASTE OR DELETE CELLS INCLUDED IN THE ASSIGNMENT.**


## How to complete assignments

Whenever you see:

```
# YOUR CODE HERE
raise NotImplementedError()
```

You need to **replace (by which we mean _delete_) these lines of code with code that answers the questions** and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

You should write the answer to the questions in those cells (the ones with `# YOUR CODE HERE`), but you can also add extra cells to explore / investigate things if you need / want to. 

Any cell with `assert` statements in it is a test cell. You should not try to change or delete these cells. Note that there might be more than one assert that tests a particular question. 

If a test does fail, reading the error that is printed out should let you know which test failed, which may be useful for fixing it.

Note that some cells, including the test cells, may be read only, which means they won't let you edit them. If you cannot edit a cell - that is normal, and you shouldn't need to edit that cell.


## Tips & Tricks

The following are a couple tips & tricks that may help you if you get stuck on anything.

#### Printing Variables
You can (and should) print and check variables as you go. This allows you to check what values they hold, and fix things if anything unexpected happens.

#### Restarting the Kernel
- If you run cells out of order, you can end up overwriting things in your namespace. 
- If things seem to go weird, a good first step is to restart the kernel, which you can do from the kernel menu above.
- Even if everything seems to be working, it's a nice check to 'Restart & Run All', to make sure everything runs properly in order.

```python
# basic plotting imports
%matplotlib inline
%config InlineBackend.figure_format = 'retina' # high res plotting

import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
sns.set(rc={'figure.figsize':(12,9)})
sns.set_style('whitegrid')
sns.set_style("whitegrid", {'axes.grid' : False})
```

# Part 1: Correlation basics

This section serves as a reminder for the basics of how correlation works. A Pearson correlation is a normalized measure of the covariance between two variables (features); more simply, it quantifies how linear the relationship is between the two variables.

Below, we will walk you through the powers, pitfalls, and idiosyncrasies of correlations.


## Question 1: Random data

To begin, use numpy.random.normal to create two normally distrubted random variables: <code>random_data1 and random_data2</code>

<code>random_data1</code> should have mean=1.0, std=1.0
<code>random_data2</code> should have mean=6.0, std=2.5

Both should have 1000 samples.


<!-- grade_id: cell-2c1f5c27f99bc71e, points: 0, grade: false, solution: true -->
```python
import numpy as np

number_of_samples = 1000

### BEGIN SOLUTION
# make two large, random variables
random_data1 = np.random.normal(1., 1., number_of_samples)
random_data2 = np.random.normal(6., 2.5, number_of_samples)
### END SOLUTION
```

<!-- grade_id: cell-e80e9bb810226a8f, points: 10, grade: true, solution: false -->
```python
# Hidden tests (worth 10 points)
### BEGIN HIDDEN TESTS
number_of_samples = 1000
assert len(random_data1) & len(random_data2) == number_of_samples
assert np.isclose(np.mean(random_data1), np.mean(np.random.normal(1., 1., number_of_samples)), atol=0.3)
assert np.isclose(np.mean(random_data2), np.mean(np.random.normal(6., 2.5, number_of_samples)), atol=0.3)
assert np.isclose(np.std(random_data1), np.std(np.random.normal(1., 1., number_of_samples)), atol=0.3)
assert np.isclose(np.std(random_data2), np.std(np.random.normal(6., 2.5, number_of_samples)), atol=0.3)
### END HIDDEN TESTS
```

## Q2: Random data correlation significance

Now let's look at the Pearson correlation between the two random datasets.

But first, given that these are two random variables, will they be correlated?

* `A`: They will <em>definitely</em> be significantly correlated (<em>p</em> < 0.05)
* `B`: They will <em>definitely not</em> be significantly correlated (<em>p</em> > 0.05) 
* `C`: The correlation between them cannot be assessed
* `D`: The correlation will probably be non-significant (<em>p</em> > 0.05) but will sometimes be significant (<em>p</em> < 0.05)

Write your answer below as a new variable, <code>q2_answer</code>. So if the answer was a hypothetical option <code>E</code>, you would write:

<code>q2_answer = 'E'</code>

<!-- grade_id: cell-0a742e3f3895f1fb, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q2_answer = 'D'
### END SOLUTION
```

<!-- grade_id: cell-70b4310f5dd968f6, points: 5, grade: true, solution: false -->
```python
# Tests for Q2 (worth 5 points)
assert isinstance(q2_answer, str)

### BEGIN HIDDEN TESTS
assert (q2_answer == 'D') or (q2_answer == 'd')
### END HIDDEN TESTS
```

## Q3: Correlating random data

Now let's look at the actual Pearson correlation.

To do this, call <code>scipy.stats.pearsonr</code> and store the Pearson correlation coefficient as <code>r_random</code> and the associated <em>p</em>-value as <code>p_random</code>.

Then plot the relationship with <code>random_data1</code> as the x-axis and <code>random_data2</code> as the y-axis. Set the opacity (alpha) of the points equal to 0.5, by defining a new variable, <code>alpha_scatter</code> and calling that in the scatterplot function.

<!-- grade_id: cell-d446a4f129f5e0f9, points: 0, grade: false, solution: true -->
```python
import scipy as sp

### BEGIN SOLUTION
alpha_scatter = 0.5
r_random, p_random = sp.stats.pearsonr(random_data1, random_data2)

# plt.scatter(random_data1, random_data2, alpha = alpha_scatter)
# plt.show()
### END SOLUTION
```

<!-- grade_id: cell-e6a9885757a4a4fe, points: 5, grade: true, solution: false -->
```python
# Hidden tests (worth 5 points)
### BEGIN HIDDEN TESTS
assert (r_random >= -1.) & (r_random <= 1.)
assert (p_random >= 0.) & (p_random <= 1.)
assert (alpha_scatter == 0.5)
### END HIDDEN TESTS
```

## Q4: Making a linear correlation

Okay, so  we know that a Pearson correlation quatifies how linear the relationship is between two variables (features).

Now, let's create a new variable, <code>correlated_data1</code>, that is explicitly a linear function of <code>random_data1</code>, plus some noise.

To do this, set <code>correlated_data1</code> equal to <code>random_data1</code> plus 20., and add normally-distributed noise of mean=0., std=0.5

Calculate the Pearson correlation coefficient between them as <code>r_corr</code> and the associated <em>p</em>-value as <code>p_corr</code>.

Plot the results using the same opacity as above.

<!-- grade_id: cell-1f056a5ab429c813, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
correlated_data1 = random_data1 + 20.
noise = np.random.normal(0., 0.5, np.size(correlated_data1))
correlated_data1 = correlated_data1 + noise

r_corr, p_corr = sp.stats.pearsonr(random_data1, correlated_data1)

# plt.scatter(random_data1, correlated_data1, alpha = alpha_scatter)
# plt.show()
### END SOLUTION
```

<!-- grade_id: cell-2206a4642f630d2f, points: 5, grade: true, solution: false -->
```python
# Hidden tests (worth 5 points)
### BEGIN HIDDEN TESTS
assert np.isclose(np.mean(correlated_data1-random_data1), 20, atol=1.)
assert np.isclose(np.std(correlated_data1-random_data1), 0.5, atol=0.2)

assert (r_corr >= 0.5) & (r_corr <= 1.)
assert (p_corr >= 0.) & (p_corr <= 0.05)

### END HIDDEN TESTS
```

# Part 2: Correlation cautions


## Q5: Assessing correlation significance

Now we're going to switch gears and give you four sets of data:

* <code>(x1, y1)</code>
* <code>(x2, y2)</code>
* <code>(x3, y3)</code>
* <code>(x4, y4)</code>

Calculate the means and variances for each of the eight variables.

To do this, make a list of the <code>x</code> variables, called <code>xs</code>; do the same for the <code>y</code> variables, called <code>ys</code>.

Then create a variable called <code>x_stats</code> that has two arrays: the means for each <code>x</code> variable, and the vars for each <code>x</code> variable; do the same for the <code>y</code> variables, called <code>y_stats</code>

Calculate the Pearson correlation coefficients and the associated <em>p</em>-values. Store the correlation coefficients as an array called <code>rs</code>, and the <em>p</em>-values as an array called <code>ps</code>.


<!-- grade_id: cell-7acaaf57f7bb9d97, points: 0, grade: false, solution: true -->
```python
# Given arrays (don't change these!)
x1 = np.array([10, 8, 13, 9, 11, 14, 6, 4, 12, 7, 5])
y1 = np.array([8.04, 6.95, 7.58, 8.81, 8.33, 9.96, 7.24, 4.26, 10.84, 4.82, 5.68])
x2 = np.array([10, 8, 13, 9, 11, 14, 6, 4, 12, 7, 5])
y2 = np.array([9.14, 8.14, 8.74, 8.77, 9.26, 8.10, 6.13, 3.10, 9.13, 7.26, 4.74])
x3 = np.array([10, 8, 13, 9, 11, 14, 6, 4, 12, 7, 5])
y3 = np.array([7.46, 6.77, 12.74, 7.11, 7.81, 8.84, 6.08, 5.39, 8.15, 6.42, 5.73])
x4 = np.array([8, 8, 8, 8, 8, 8, 8, 19, 8, 8, 8])
y4 = np.array([6.58, 5.76, 7.71, 8.84, 8.47, 7.04, 5.25, 12.50, 5.56, 7.91, 6.89])

### BEGIN SOLUTION
xs = np.array([x1, x2, x3, x4])
x_stats = np.mean(xs, axis=1), np.var(xs, axis=1)

ys = np.array([y1, y2, y3, y4])
y_stats = np.mean(ys, axis=1), np.var(ys, axis=1)

total_variables = xs.shape[0]
rs = np.zeros(total_variables)
ps = np.zeros(total_variables)
for i in range(total_variables):
    rs[i], ps[i] = sp.stats.pearsonr(xs[i], ys[i])

### END SOLUTION
```

```python
x_stats[0]
```

Output:
```
array([9., 9., 9., 9.])
```

<!-- grade_id: cell-91d11144709c1a19, points: 10, grade: true, solution: false -->
```python
# Hidden tests (worth 10 points)
### BEGIN HIDDEN TESTS
assert np.shape(xs) == (4, 11)
assert np.shape(ys) == (4, 11)

assert np.shape(x_stats) == (2, 4)
assert np.shape(y_stats) == (2, 4)

assert np.shape(rs) == (4,)
assert np.shape(ps) == (4,)

assert np.all(x_stats[1]-10 == [0., 0., 0., 0.])
assert np.all(np.isclose(y_stats[1]-3.75, 0., atol=0.01))
### END HIDDEN TESTS
```

## Q6: Interpreting these results

Okay, so given the means, variances, and Pearson correlation coefficients (collectively, "summary statistics") for the four above datasets, what inference do you make about them?

* `A`: The summary statistics are all very close to the same
* `B`: The means and variances are close to the same, but the Pearson correlations are different
* `C`: The means and variances are very different, but the Pearson correlations are very similar
* `D`: All the summary statistics are very different from one another

Write your answer below as a new variable, <code>q6_answer</code>. So if the answer was a hypothetical option <code>E</code>, you would write:

<code>q6_answer = 'E'</code> 

<!-- grade_id: cell-ee491b37d9d163d4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q6_answer = 'A'
### END SOLUTION
```

<!-- grade_id: cell-0c4b0bd614377625, points: 5, grade: true, solution: false -->
```python
# Tests for Q6 (worth 5 points)
assert isinstance(q6_answer, str)

### BEGIN HIDDEN TESTS
assert (q6_answer == 'A') or (q6_answer == 'a')
### END HIDDEN TESTS
```

## Q7: Summary statistics: a cautionary tale

Now plot the four <code>(x, y)</code> pairs. What inference do you make about them, now?

* `A`: The <code>(x, y)</code> pairs all look very similar
* `B`: The <code>(x, y)</code> pairs all look linear
* `C`: The <code>(x, y)</code> pairs all look very different
* `D`: The <code>(x, y)</code> pairs all look quadratic

Write your answer below as a new variable, <code>q7_answer</code>. So if the answer was a hypothetical option <code>E</code>, you would write:

<code>q7_answer = 'E'</code> 

<!-- grade_id: cell-b0242660c36fe0b9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
q7_answer = 'C'

# fig, axs = plt.subplots(2, 2)
# axs[0, 0].scatter(x1, y1, alpha = 0.8)
# axs[0, 1].scatter(x2, y2, alpha = 0.8)
# axs[1, 0].scatter(x3, y3, alpha = 0.8)
# axs[1, 1].scatter(x4, y4, alpha = 0.8)
# plt.show()
### END SOLUTION
```

<!-- grade_id: cell-041db26acd669a14, points: 5, grade: true, solution: false -->
```python
# Tests for Q7 (worth 5 points)
assert isinstance(q7_answer, str)

### BEGIN HIDDEN TESTS
assert (q7_answer == 'C') or (q7_answer == 'c')
### END HIDDEN TESTS
```

## Q8: Alternative correlation approaches

Having plot the data, you may have noticed certain features that might lead the Pearson correlation to return results that don't match with your intuitive relationship between the variables.

Let's figure out some ways of addressing that.

We're going to create two variables. The first, <code>x_axis</code>, is supplied below. It is simply 10 numbers, evenly spaced, between 1 and 100.

The second should be a function of <code>x_axis</code>. Specifically, create a new variable, <code>nonlinear_data</code> that is <code>x_axis</code> divided by <code>x_axis**2</code>. Plot the results.

Last, calculate the Pearson correlation and associated <em>p</em>-value as <code>r_nonlinear</code> and <code>p_nonlinear</code>, respectively.

<!-- grade_id: cell-ed2a867c0bd42af4, points: 0, grade: false, solution: true -->
```python
x_axis = np.linspace(1, 100, 10)

### BEGIN SOLUTION
nonlinear_data = x_axis / (x_axis**2)
r_nonlinear, p_nonlinear = sp.stats.pearsonr(x_axis, nonlinear_data)

# plt.scatter(x_axis, nonlinear_data)
# plt.show()
### END SOLUTION
```

<!-- grade_id: cell-e9342f30ae7182e3, points: 5, grade: true, solution: false -->
```python
# Hidden tests (worth 5 points)
### BEGIN HIDDEN TESTS
assert np.all(nonlinear_data == (x_axis / (x_axis**2)))
assert np.isclose(r_nonlinear, -0.5721, atol=0.01)
assert np.isclose(p_nonlinear, 0.08396, atol=0.01)
### END HIDDEN TESTS
```

## Q9: Summary statistics: a cautionary tale

From the plot above (and from the equation), you can see that <code>nonlinear_data</code> is monotonically related to <code>x_axis</code>, such that each successive value of <code>nonlinear_data</code> is smaller than the previous. But the <em>p</em>-value from the Pearson correlation significance test is non-significant (<em>p</em> > 0.05).

A non-parameteric alternative to the Pearson correlation is the Spearman correlation (<code>scipy.stats.spearmanr</code>). This correlation approach examines the monotonicity of the relationship between the two variables. That is, rather than asking if they're linearly related, it asks how much one variable regularly increases (or decreases) in relation to the other.

For this question, calculate the Spearman correlation between <code>x_axis</code> and <code>nonlinear_data</code>, and the associated <em>p</em>-value, as <code>rho_nonlinear</code> and <code>pho_nonlinear</code>, respectively.

Do the same for all <code>(x, y)</code> pairs above. Store the correlation coefficients as an array called <code>rho_spearman</code>, and the <em>p</em>-values as an array called <code>p_spearman</code>.

Note how the Spearman correlation coefficiencts ($\rho$) and <em>p</em>-values differ from the Pearson ones.

<!-- grade_id: cell-f0bf1f90dfa785ee, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
rho_nonlinear, pho_nonlinear = sp.stats.spearmanr(x_axis, nonlinear_data)

total_variables = xs.shape[0]
rho_spearman = np.zeros(total_variables)
p_spearman = np.zeros(total_variables)
for i in range(total_variables):
    rho_spearman[i], p_spearman[i] = sp.stats.spearmanr(xs[i], ys[i])

# print(rho_nonlinear, pho_nonlinear)
# print(rho_spearman, p_spearman)
### END SOLUTION
```

<!-- grade_id: cell-dc76cfae32631e09, points: 10, grade: true, solution: false -->
```python
# Hidden tests (worth 10 points)
### BEGIN HIDDEN TESTS
assert rho_nonlinear < -0.99
assert pho_nonlinear < 10e-40

assert np.shape(rho_spearman) == (4,)
assert np.shape(p_spearman) == (4,)

assert np.all(np.isclose((rho_spearman - [0.81818, 0.69090, 0.9909, 0.5]), 0., atol=0.01))
assert np.all(p_spearman < 0.2)
### END HIDDEN TESTS
```

## Q10: Model testing: Nonlinear models

Okay so now what do we do with this? Some of those <code>(x, y)</code> pairs above look like they have clear relationships (nonlinear, or linear with outliers). How do we handle this kind of thing?

Here we're going to do a little basic model testing using <code>scipy.optimize.curve_fit</code>. This function lets us define any function of the form:

$$ y = f(x) $$

In the cell below I show how you define a linear function:

$$ y = ax + b $$

pass that into <code>scipy.optimize.curve_fit</code>, and plot the fit.

You job is to write another function, <code>quadratic_func</code>:

$$ y = ax^2 + bx + c $$

and calculate both the linear and quadratic $R^2$ fits for all of the <code>(x, y)</code> pairs. I've provided the $R^2$ function, <code>fit_r2</code>, for you.

Store the linear $R^2$ fits as an array called <code>linear_r2</code> and the quadratic as <code>quadratic_r2</code>.

<!-- grade_id: cell-dd692b604a366d35, points: 0, grade: false, solution: true -->
```python
from scipy.optimize import curve_fit

def linear_func(x, a, b):
    return (a * x) + b

def fit_r2(ydata, yfit):
    residuals = ydata - yfit # difference between data and fit
    ss_res = np.sum(residuals**2) # sum of the squares (ss) of the residuals
    ss_tot = np.sum((ydata - np.mean(ydata))**2) # total sum of squares
    r2 = 1 - (ss_res / ss_tot) # r-squared    
    return r2

xdata = xs[0,:]
ydata = ys[0,:]
popt, _ = curve_fit(linear_func, xdata, ydata)
yfit = linear_func(xdata, *popt)

# Because it can be non-intuitive, the below code will plot
# your data and the fits, so you can test your code out
#
# indices = np.argsort(xdata)
# yfit_sorted = yfit[indices]
# xfit_sorted = xdata[indices]
# plt.plot(xdata, ydata, 'o', label='data')
# plt.plot(xfit_sorted, yfit_sorted, '--', label='fit')
# plt.legend()
# plt.show()

### BEGIN SOLUTION
def quadratic_func(x, a, b, c):
    return (a * (x**2)) + (b * x) + c

total_variables = xs.shape[0]
linear_r2 = np.zeros(total_variables)
quadratic_r2 = np.zeros(total_variables)
for i in range(total_variables):
    xdata = xs[i,:]
    ydata = ys[i,:]
    yfit = []
    popt = []
    
    popt, _ = curve_fit(linear_func, xdata, ydata)
    yfit = linear_func(xdata, *popt)
    linear_r2[i] = fit_r2(ydata, yfit)
    
    popt, _ = curve_fit(quadratic_func, xdata, ydata)
    yfit = quadratic_func(xdata, *popt)
    quadratic_r2[i] = fit_r2(ydata, yfit)
### END SOLUTION
```

<!-- grade_id: cell-e2c7289b2e098aad, points: 15, grade: true, solution: false -->
```python
# Hidden tests (worth 15 points)
### BEGIN HIDDEN TESTS
assert quadratic_func(5, 3, 2, 1) == 86
assert np.shape(linear_r2) == (4,)
assert np.shape(quadratic_r2) == (4,)
assert np.all(np.isclose((linear_r2 - [0.666, 0.666, 0.666, 0.666]), 0., atol=0.001))
assert np.all(np.isclose((quadratic_r2 - [0.687, 0.999, 0.684, 0.666]), 0., atol=0.001))
### END HIDDEN TESTS
```

## Q11: Model testing: Resampling

You'll see that the quadratic fit dramatically improves the fit quality for one of the <code>(x, y)</code> pairs, but doesn't do much for the others. So another way to test model fits is by using a broad class of resampling methods. In this approach, instead of fitting all of your data, you select a random segemnt of your data, fit those points and store the fit results, draw another random sample, fit those points and store the results, and so on.

Here, I want you to calculate the $R^2$ linear and quadratic fits for <em>all but one</em> data point for each of the four <code>(x, y)</code> pairs. That is, for each <code>(x, y)</code> pairs, fit everything except the zeroth element (so <code>xs[:, 1:]</code>). Then iterate by fitting all points except the first element, and then iterate again for all but the second, and so on.

Store the resampled linear $R^2$ fits as a <code>(4, 11)</code> array called <code>resampled_linear_r2</code> and the quadratic as an array of the same size called <code>resampled_quadratic_r2</code>.

<!-- grade_id: cell-38ce5f100b5276f4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
total_variables = xs.shape[0]
total_resamples = xs.shape[1]
resampled_linear_r2 = np.zeros([total_variables, total_resamples])
resampled_quadratic_r2 = np.zeros([total_variables, total_resamples])

indices = np.arange(total_resamples)

for i in range(total_variables):
    for j in range(total_resamples):
        xdata = xs[i, np.delete(indices, j)]
        ydata = ys[i, np.delete(indices, j)]
        yfit = []
        popt = []

        popt, _ = curve_fit(linear_func, xdata, ydata)
        yfit = linear_func(xdata, *popt)
        resampled_linear_r2[i, j] = fit_r2(ydata, yfit)

        popt, _ = curve_fit(quadratic_func, xdata, ydata)
        yfit = quadratic_func(xdata, *popt)
        resampled_quadratic_r2[i, j] = fit_r2(ydata, yfit)
### END SOLUTION
```

<!-- grade_id: cell-acd0e5c7da168d4c, points: 10, grade: true, solution: false -->
```python
# Hidden tests (worth 10 points)
### BEGIN HIDDEN TESTS
assert np.shape(resampled_linear_r2) == (4,11)
assert np.shape(resampled_quadratic_r2) == (4,11)
assert np.all(np.isclose((np.mean(resampled_linear_r2, 1) - [0.666, 0.667, 0.694, 0.628]), 0., atol=0.001))
assert np.all(np.isclose((np.mean(resampled_quadratic_r2, 1) - [0.689, 0.999, 0.718, 0.628]), 0., atol=0.001))
### END HIDDEN TESTS
```

## Final takeaways

Look at the distribution of possible correlation coefficients that your data can take. For example, one of the <code>(x, y)</code> looks like a clear linear fit with a single large outlier; in that case, the distribution of correlation coefficients all look similar, except for a single case where the $R^2$ is nearly equal to 1.0. Similarly, another of the pairs looks like a perfect quadratic fit in every single resampling scenario.

What you do, as a scientist, with this information, is part of the "art" of the data analysis decision making process!

```python

```

```python

```